```python
# Initialize Otter
import otter
grader = otter.Notebook("p7.ipynb")
```

```python
import p7_test
import csv
import math
```

```python
# PLEASE FILL IN THE DETAILS
# Enter none if you don't have a project partner

# project: p7
# submitter: abrajanrojas
# partner: tlxiong
```

# Project 7: Drinking Water Accessibility

## Learning Objectives:

In this project you will demonstrate how to:

- Write programs to interpret data present in csv files,
- Use lists and dictionaries effectively to manage data,
- **Develop good coding styling habits (points may be deducted for bad coding styles)**

## Testing your code:

Along with this notebook, you must have downloaded the file `p7_test.py`. If you are curious about how we test your code, you can explore this file, and specifically the value of the variable `expected_json`, to understand the expected answers to the questions. You can have a look at [p2](https://git.doit.wisc.edu/cdis/cs/courses/cs220/cs220-f22-projects/-/tree/main/p2/p2.ipynb) if you have forgotten how to read the outputs of the `grader.check` function calls.

**Please go through [lab-p7](https://git.doit.wisc.edu/cdis/cs/courses/cs220/cs220-f22-projects/-/tree/main/lab-p7) before starting this project.** The lab introduces some useful techniques necessary for this project.

## Project Description:

Universal access to safe drinking water is a fundamental need and human right. Securing access for all would go a long way in reducing illness and death, especially among children. "Safely managed" drinking water services represent an ambitious new rung on the ladder used to track progress on drinking water. Since 2000, 2 billion people have gained access to safely managed services (i.e., accessible on-premises, available when needed, and free from contamination). In 2020, 5.8 billion people used safely managed services and a further 2 billion people used basic services. However, 771 million people still lacked even a basic level of service, including 282 million who used a “limited” water service (source from which water collection time exceeds 30 minutes), 367 million who used unimproved sources and 122 million who still collected drinking water directly from rivers, lakes, and other surface water sources. The data reveal pronounced disparities, with the poorest and those living in rural areas least likely to use a basic service. In most countries, the burden of water collection continues to fall mainly to women and girls.

[The Unicef website](https://data.unicef.org/) states that "consistent, credible data about children’s situations are critical to the improvement of their lives – and indispensable to realizing the rights of every child." Data Scientists will play an important role in reaching this goal.

For this project, you'll be analyzing data drawn from multiple sources. Our data is primarily drawn from the report titled ["Progress on Household Drinking Water, Sanitation and Hygiene"](https://washdata.org/sites/default/files/2021-07/jmp-2021-wash-households.pdf) data published by the Unicef/WHO Joint Monitoring Programme for Water Supply, Sanitation and Hygiene (2021). The original dataset can be found [here](https://data.unicef.org/topic/water-and-sanitation/drinking-water/) if you are interested in exploring the dataset yourself. Our dataset is further augmented by data from The World Bank on the [income levels of each country](https://datatopics.worldbank.org/world-development-indicators/the-world-by-income-and-region.html).

## Dataset:

The JMP report defines *people who have access to an [improved source of water](https://www.cdc.gov/healthywater/global/assessing.html#ImprovedDrinking) within 30 minutes round trip collection time* as having [at least basic access](https://www.cdc.gov/healthywater/global/assessing.html#DrinkingWaterSources) to water. For this project, we will focus on the **percentage of population** of each country who had **at least basic** water supply in the years **2015** and **2020**. Open `water_accessibility.csv` with Microsoft Excel or some other Spreadsheet viewer and look at the list of countries in the dataset. Data for each country appears twice, one row for the year *2015* and the other row for year *2020*. Countries which had incomplete data have been **omitted** from the dataset, and we will **ignore** those countries in this project. You do **not** have to deal with any **missing data** in the dataset.

The data shows:
- `country_code` : the unique country code that consists of three alphabet letters
- `country_name` : the name of the country
- `region` : the geographical location of the country (does not equal to its corresponding continents, but follows the administrative groupings from [The World Bank](https://datahelpdesk.worldbank.org/knowledgebase/articles/906519-world-bank-country-and-lending-groups))
- `year` : the year in which it was subject to data collection
- `income_level` : the classification of income level based on GNI per capita in US dollars ([The World Bank Atlas Method](https://datahelpdesk.worldbank.org/knowledgebase/articles/378834-how-does-the-world-bank-classify-countries))
- `pop` : population of the country in a specific year (in thousands)
- `urban_percent` : the percentage of population in a given country that is urban
- `national_alb` : the percentage of a country's population that has access to at least basic water supply
- `urban_alb` : the percentage of a country's urban population that has access to at least basic water supply

## Project Requirements:

You **may not** hardcode indices in your code, unless the question explicitly . If you open your `.csv` file with Excel, manually count through the rows and use this number to loop through the dataset, this is also considered as hardcoding. We'll **manually deduct** points from your autograder score on Gradescope during code review.

**Store** your final answer for each question in the **variable specified for each question**. This step is important because Otter grades your work by comparing the value of this variable against the correct answer.

For some of the questions, we'll ask you to write (then use) a function to compute the answer. If you compute the answer **without** creating the function we ask you to write, we'll **manually deduct** points from your autograder score on Gradescope, even if the way you did it produced the correct answer. 

Required Functions:
- cell
- get_col_dict

In this project, you will also be required to define certain **data structures**. If you do not create these data structures exactly as specified, we'll **manually deduct** points from your autograder score on Gradescope, even if the way you did it produced the correct answer.

Required Data Structures:
- dict_2015
- dict_2020
- rural_non_alb_bin_2015_dict
- rural_non_alb_bin_2020_dict
    
Students are only allowed to use Python commands and concepts that have been taught in the course prior to the release of p7. Therefore, **you should not use the pandas module**.  We will **manually deduct** points from your autograder score on Gradescope otherwise.

For more details on what will cause you to lose points during code review and specific requirements, please take a look at the [Grading rubric](https://git.doit.wisc.edu/cdis/cs/courses/cs220/cs220-f22-projects/-/blob/main/p7/rubric.md).

## Incremental Coding and Testing:

You should always strive to do incremental coding. **Incremental coding enables you to avoid challenging bugs.** Always write a few lines of code and then test those lines of code, before proceeding to write further code. You can call the `print` function to test intermediate step outputs.

We also recommend you do incremental testing: make sure to run the local tests as soon as you are done with a question. This will ensure that you haven't made a big mistake that might potentially impact the rest of your project solution. Please refrain from making multiple submissions on Gradescope for testing individual questions' answers. Instead use the local tests, to test your solution on your laptop.

That said, it is **important** that you check the Gradescope test results as soon as you submit your project on Gradescope. Test results on Gradescope are typically available somewhere between 2 to 10 minutes after the submission.

Also, remember to check with the [p7 rubric](https://git.doit.wisc.edu/cdis/cs/courses/cs220/cs220-f22-projects/-/blob/main/p7/rubric.md) to verify that you will not be losing any points during manual review.

## Project Questions and Functions:

```python
# it is considered a good coding practice to place all import statements at the top of the notebook
# please place all your import statements in this cell if you need to import any more modules for this project

```

First, read the data stored in `water_accessibility.csv` and store it in a suitable variable. You can follow the same procedure you used in p6.

```python
# read the data stored in water_accessibility.csv
def process_csv(filename):
    example_file = open(filename, encoding="utf-8")
    example_reader = csv.reader(example_file)
    example_data = list(example_reader)
    example_file.close()
    return example_data
```

```python
# call process_csv here to read 'water_accessibility.csv'
csv_data = process_csv('water_accessibility.csv')

# split the header and data into appropriate variables
csv_header = csv_data[0]
csv_rows = csv_data[1:]
```

### Function 1: `cell(row_idx, col_name)` 

This function must take in a row index, `row_idx` and a column name, `col_name` as its inputs, and return the value in `water_accessibility.csv` stored there. There is **no missing data** in this dataset.

You **must** define the variables `csv_header` and `csv_rows` as in lab-p7, and you **must** copy/paste your `cell` function from lab-p7.

**Important:** You **must** only use the `cell` function to extract data from the dataset. If you extract any data without explicitly using this function, you will **lose points** during manual review.

```python
# define the cell function here
def cell(row_idx, col_name):
    col_idx = csv_header.index(col_name)
    val = csv_rows[row_idx][col_idx]
    if val == "":
        return None
    elif col_name == 'country_code':
        val = str(val)
    elif col_name == 'country_name':
        val = str(val)
    elif col_name == 'region':
        val = str(val)
    elif col_name == 'year':
        val = int(val)
    elif col_name == 'income_level':
        val = str(val)
    elif col_name == 'pop':
        val = int(val) * 1000
    elif col_name == 'urban_percent':
        val = int(val)
    elif col_name == 'national_alb':
        val = int(val)
    elif col_name == 'urban_alb':
        val = int(val)
    return val
```

You are all set! You are now ready to start solving the questions.

**Question 1:** Which country had the highest population (`pop`) in *2020*?

```python
# compute and store the answer in the variable 'highest_pop_country', then display it

highest_pop_country = []
pop_val = 0
for idx in range(1, len(csv_data), 2):
    high_pop = cell(idx, "pop")
    country = cell(idx, "country_name")
    if pop_val < high_pop:
        country = cell(idx, "country_name")
        pop_val += high_pop
        highest_pop_country.append(country)

highest_pop_country = highest_pop_country[3]
highest_pop_country
```

Output:
```
'China'
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

**Question 2:** Which country had the highest population (`pop`) **increase** between *2015* and *2020*?

There is a **unique** country in this dataset whose population increased the most. You **do not** have to worry about ties.

**Hint:** Recall how to loop through the dataset and extract data from each year from [lab-p7](https://git.doit.wisc.edu/cdis/cs/courses/cs220/cs220-f22-projects/-/tree/main/lab-p7).

```python
# compute and store the answer in the variable 'highest_pop_inc_country', then display it

highest_pop_inc_country = ''
highest_pop_inc = 0
# you need to loop through only even indices to avoid duplicate counting of the rows
for idx in range(0, len(csv_rows), 2):
    idx_2015 = cell(idx,'pop')
    country = cell(idx, 'country_name')
    idx_2020 = cell(idx+1,'pop')
    pop_growth = idx_2020 - idx_2015
    #highest_pop_inc = max(int(pop_growth))
    
    if pop_growth > highest_pop_inc: # check if population in 2015 is less than population in 2020
        highest_pop_inc_country = country
        highest_pop_inc += pop_growth
        
        
highest_pop_inc_country
```

Output:
```
'India'
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

**Question 3:** Which country had the highest **increase** in at least basic (`national_alb`) water supply between the years of *2015* and *2020*? 

There is a **unique** country in this dataset whose `national_alb` value increased the most. You **do not** have to worry about ties.

**Hint:** Take a look at q7 in lab-p7 to see how to compute the change in `national_alb` between *2015* and *2020* for each country.

```python
highest_nat_alb_inc_country = ''
highest_nat_alb = 0

for idx in range(0, len(csv_rows), 2):
    idx_2015 = cell(idx,'national_alb')
    country = cell(idx,'country_name')
    idx_2020 = cell(idx+1, 'national_alb')
    alb_growth = idx_2020 - idx_2015
    
    if alb_growth > highest_nat_alb:
        highest_nat_alb_inc_country = country
        highest_nat_alb += alb_growth
        
highest_nat_alb_inc_country
```

Output:
```
'Afghanistan'
```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
```

**Question 4:** What was the `income_level` in *2020* of the country with the highest increase in at least basic (`national_alb`) water supply between *2015* and *2020*?

```python
# compute and store the answer in the variable 'highest_alb_inc_income_level', then display it
highest_alb_inc_income_level = ''
highest_alb_inc_income = 0
for idx in range(0,len(csv_rows),2):
    #country = cell(idx,'country_name')
    income_lvl_2020 = cell(idx+1,'income_level')
    alb_2015 = cell(idx,'national_alb')
    alb_2020 = cell(idx+1, 'national_alb')
    alb_growth = alb_2020 - alb_2015
    if alb_growth > highest_alb_inc_income: #and income_lvl_2020:
        highest_alb_inc_income_level = income_lvl_2020
        highest_alb_inc_income = alb_growth
highest_alb_inc_income_level
```

Output:
```
'Low income'
```

```python
grader.check("q4")
```

Output:
```
q4 results: All test cases passed!
```

**Question 5:** What is the **total** population (`pop`) of **all** the countries (in the dataset) in the `year` *2020*?

The `pop` column stores the population in thousands. So, you will have to multiply the population of each country by *1000* to get (an approximation of) the actual number.

**Hint:** You may modify your `cell` function, so that it automatically multiplies the values in the `pop` column by *1000* before returning.

```python
# compute and store the answer in the variable 'total_pop_2020', then display it

total_pop_2020 = 0

for idx in range(0,len(csv_rows), 2):
    idx_2020 = cell(idx+1,'year')
    population = cell(idx+1,'pop')
    if idx_2020 == 2020:
        total_pop_2020 += population
        
total_pop_2020
```

Output:
```
6861226000
```

```python
grader.check("q5")
```

Output:
```
q5 results: All test cases passed!
```

**Question 6:** What was the global **percentage** of urban population (`urban_percent`) across **all** countries (in the dataset) in the `year` *2020*?

You need to find the **total** urban population by *multiplying* the `pop` and `urban_percent` columns of each country and *adding* it up. Then you need to *divide* by the **total** population to get the percentage of urban population across all the countries.

Your output **must** be an **int**. You **must** use the `round` function to round your answer to the nearest integer.

```python
# compute and store the answer in the variable 'urban_pop_percent_2020', then display it

urban_pop_percent_2020 = ''
total_urb_pop = 0
total_pop = 0
for idx in range(0, len(csv_rows), 2):
    idx_2020 = cell(idx+1,'year')
    population = cell(idx+1,'pop')
    urban_percent = cell(idx+1,'urban_percent')
    TUP = population * urban_percent
    total_urb_pop += TUP
    total_pop += population
    div_per = total_urb_pop / total_pop
    urban_pop_percent_2020 = math.ceil(div_per)
    
urban_pop_percent_2020
```

Output:
```
55
```

```python
grader.check("q6")
```

Output:
```
q6 results: All test cases passed!
```

**Question 7:** What was the **total** population (`pop`) of countries that were in the *High income* group (`income_level`) in the `year` *2015*?

Your output **must** be an **int**.

```python
# compute and store the answer in the variable 'high_income_pop', then display it
high_income_pop = 0

for idx in range(0, len(csv_rows), 2):
    idx_2015 = cell(idx, 'year')
    income_level = cell(idx,'income_level')
    population = cell(idx,'pop')
    if idx_2015 == 2015 and income_level == 'High income':
        high_income_pop += population
        
high_income_pop
```

Output:
```
848304000
```

```python
grader.check("q7")
```

Output:
```
q7 results: All test cases passed!
```

**Question 8:** Which *income group* (`income_level`) had the **least** population (`pop`) in the `year` *2015*?

You must find the **total** population (`pop`) for each `income_level`, and find the `income_level` which has the **least** total population.

**Hint:** There are several ways to solve this problem efficiently (including using `dicts`). You can try to solve this problem using dicts if you want to. However, another approach that you might already be familiar with from p6 is to first create a *list* of all the **unique** income levels, and then loop through the entire dataset for **each** income level to find the total population of that income level, before comparing these numbers to find the income level with the least population.

```python
# compute and store the answer in the variable 'least_pop_income_group', then display it
least_pop_income_group = []
low_pop = 0
pop_total = 0

for idx in range(len(csv_rows)):
    least_pop = int(cell(idx, "pop"))
    if least_pop > 0:
        income_group = cell(idx, "income_level")
        least_pop_income_group.append(income_group)

least_pop_income_group = str(least_pop_income_group[1])
least_pop_income_group
```

Output:
```
'Low income'
```

```python
grader.check("q8")
```

Output:
```
q8 results: All test cases passed!
```

**Question 9:** Create a **list** of the names (`country_name`) of all countries in the *North America* `region` that **tied** for the **maximum** `national_alb` in *2015*.

You need to first find the **maximum** value of `national_alb` among all countries in the `region` *North America*, and then make a `list` of **all** the countries in this region having this `national_alb` value. **Multiple** countries from *North America* have the same maximum `national_alb` value, so your output **must** be a `list` of **all** those countries.

```python
# compute and store the answer in the variable 'na_max_alb_countries', then display it
na_max_alb_countries = []
for idx in range(0, len(csv_rows), 2):
    country = cell(idx, 'country_name')
    region = cell(idx, 'region')
    if cell(idx, 'region') == 'North America':
       
        na_max_alb_countries.append(country)

na_max_alb_countries
```

Output:
```
['Bermuda', 'Canada', 'United States of America']
```

```python
grader.check("q9")
```

Output:
```
q9 results: All test cases passed!
```

### Function 2: `get_col_dict(col_name, year)` 

This function should take in a column `col_name` and a `year` (*2015* or *2020*) as its inputs, and return a `dict` where each key is a `country_code` and the corresponding value is the value under the given `col_name` for the country with the said `country_code` in the given `year`.

For example, the value returned by `get_col_dict('country_name', 2015)` should be something like the following:
```python
{'AFG': 'Afghanistan',
 'ALB': 'Albania',
 'DZA': 'Algeria',
 'AND': 'Andorra',
 'AGO': 'Angola',
 'ARM': 'Armenia',
 'AUS': 'Australia',
 'AUT': 'Austria',
 'AZE': 'Azerbaijan',
 'BGD': 'Bangladesh',
 ...
}
```

and the value returned by `get_col_dict('pop', 2020)` should be something like the following:
```python
{'AFG': 38928000,
 'ALB': 2878000,
 'DZA': 43851000,
 'AND': 77000,
 'AGO': 32866000,
 'ARM': 2963000,
 'AUS': 25500000,
 'AUT': 9006000,
 'AZE': 10139000,
 'BGD': 164689000,
 ...
}
```

Start with the following code snippet and complete the function.

```python
# replace the ... with your code
def get_col_dict(col_name, year):
    col_dict = {}
    
    if year == 2015:
        for idx in range (0, len(csv_rows), 2):
            country = cell(idx, 'country_code')
            if col_name in ['pop', 'year', 'urban_percent', 'national_alb', 'urban_alb']:               
                col_dict[country] = int(cell(idx, col_name))
            else:
                col_dict[country] = cell(idx, col_name)
                    
    elif year == 2020:
        for idx in range (1, len(csv_rows), 2):
            country = cell(idx, 'country_code')
            if col_name in ['pop', 'year', 'urban_percent', 'national_alb', 'urban_alb']:               
                col_dict[country] = int(cell(idx, col_name))
            else:
                col_dict[country] = cell(idx, col_name)
    return col_dict
```

After you define the function `get_col_dict`, run the following two cells to test whether it works.

```python
get_col_dict_test1 = get_col_dict('region', 2020)
get_col_dict_test2 = get_col_dict('national_alb', 2015)
get_col_dict_test3 = get_col_dict('pop', 2020)
```

```python
grader.check("get_col_dict")
```

Output:
```
get_col_dict results: All test cases passed!
```

### Data Structures 1: `dict_2015`

You must now create a data structure named `dict_2015`. This data structure must be a **dict**. Each key must be a `country_code`, and the corresponding value must be another **dict**. As for the inner dictionary, the keys must be the various column names, and the values must be the values under the column name for `country_code` in the `year` *2015*.

The keys for each of the *inner* dictionary are the column names:
- `'country_name'`
- `'region'`
- `'income_level'`
- `'year'`
- `'pop'`
- `'urban_percent'`
- `'national_alb'`
- `'urban_alb'`

You are **allowed** to *hardcode* the **names** of all these columns (i.e., the keys of the *inner* dictionaries).

The data structure `dict_2015` should look something like this:
```python
{'AFG': {'country_name': 'Afghanistan',
  'region': 'South Asia',
  'income_level': 'Low income',
  'year': 2015,
  'pop': 34414000,
  'urban_percent': 25,
  'national_alb': 61,
  'urban_alb': 87},
 'ALB': {'country_name': 'Albania',
  'region': 'Europe & Central Asia',
  'income_level': 'Upper middle income',
  'year': 2015,
  'pop': 2891000,
  'urban_percent': 57,
  'national_alb': 93,
  'urban_alb': 95},
  ...
}
```

```python
# define the variable 'dict_2015' here as described above
# you may display the variable for testing purposes while you define it,
# BUT you MUST remove the line displaying 'dict_2015' before submission as the output will be too large to display

# initialize as an empty dictionary
dict_2015 = {}
for idx in range(0, len(csv_rows), 2):
    country_code = cell(idx, 'country_code')
    region = cell(idx, 'region')

    #country_dict = get_col_dict('country_name', 2015)
    #region_dict = get_col_dict('region', 2015)
    #year_dict = get_col_dict('year', 2015)
    #income_dict = get_col_dict('income_level', 2015)
    #year_dict = get_col_dict('year', 2015)
    #pop_dict = get_col_dict('pop', 2015)
    #urban_percent_dict = get_col_dict('urban_percent', 2015)
    #national_alb_dict = get_col_dict('national_alb', 2015)
    #urban_alb_dict = get_col_dict('urban_alb', 2015)

# call get_col_dict for other columns
    dict_2015[country_code] = {}
    #dict_2015[country_code][region] = {}
    dict_2015[country_code]['country_name'] = cell(idx, 'country_name')
    dict_2015[country_code]['region'] = cell(idx, 'region')
    dict_2015[country_code]['income_level'] = cell(idx, 'income_level')
    dict_2015[country_code]['year'] = int(cell(idx, 'year'))
    dict_2015[country_code]['pop'] = cell(idx, 'pop')
    dict_2015[country_code]['urban_percent'] = int(cell(idx, 'urban_percent'))
    dict_2015[country_code]['national_alb'] = int(cell(idx, 'national_alb'))
    dict_2015[country_code]['urban_alb'] = int(cell(idx, 'urban_alb')) #cell(idx, "country_code")
# add data from all these dicts to dict_2015
dict_2015
```

Output:
```
{'AFG': {'country_name': 'Afghanistan',
  'region': 'South Asia',
  'income_level': 'Low income',
  'year': 2015,
  'pop': 34414000,
  'urban_percent': 25,
  'national_alb': 61,
  'urban_alb': 87},
 'ALB': {'country_name': 'Albania',
  'region': 'Europe & Central Asia',
  'income_level': 'Upper middle income',
  'year': 2015,
  'pop': 2891000,
  'urban_percent': 57,
  'national_alb': 93,
  'urban_alb': 95},
 'DZA': {'country_name': 'Algeria',
  'region': 'Middle East & North Africa',
  'income_level': 'Upper middle income',
  'year': 2015,
  'pop': 39728000,
  'urban_percent': 71,
  'national_alb': 93,
  'urban_alb': 95},
 'AND': {'country_name': 'Andorra',
  'region': 'Europe & Central Asia',
  'income_level': 'High income',
  'year': 2015,
  'pop': 78000,
  'urban_percent': 88,
  'national_alb': 100,
  'urban_alb': 100},
 'AGO': {'country_name': 'Angola',
  'region': 'Sub-Saharan Africa',
  'income_level': 'Upper middle income',
  'year': 2015,
  'pop': 27884000,
  'urban_percent': 6
```

After you define the data structure `dict_2015`, run the following cell to test whether you have defined it properly.

```python
grader.check("dict_2015")
```

Output:
```
dict_2015 results: All test cases passed!
```

### Data Structures 2: `dict_2020`

You must now create a data structure named `dict_2020`. This data structure must be a **dict**. Each key must be a `country_code`, and the corresponding value must be another **dict**. As for the inner dictionary, the keys must be the various column names, and the values must be the values under the column name for `country_code` in the `year` *2020*.

The keys for each of the *inner* dictionary are the column names:
- `'country_name'`
- `'region'`
- `'income_level'`
- `'year'`
- `'pop'`
- `'urban_percent'`
- `'national_alb'`
- `'urban_alb'`

You are **allowed** to *hardcode* the **names** of all these columns (i.e., the keys of the *inner* dictionaries).

```python
# define the variable 'dict_2020' here as described above
# you may display the variable for testing purposes while you define it,
# BUT you MUST remove the line displaying 'dict_2020' before submission as the output will be too large to display
# initialize as an empty dictionary
dict_2020 = {}
for idx in range(len(csv_rows)):
    country_code = cell(idx, 'country_code')
    region = cell(idx, 'region')

    dict_2020[country_code] = {}
    dict_2020[country_code]['country_name'] = cell(idx, 'country_name')
    dict_2020[country_code]['region'] = cell(idx, 'region')
    dict_2020[country_code]['income_level'] = cell(idx, 'income_level')
    dict_2020[country_code]['year'] = int(cell(idx, 'year'))
    dict_2020[country_code]['pop'] = cell(idx, 'pop')
    dict_2020[country_code]['urban_percent'] = int(cell(idx, 'urban_percent'))
    dict_2020[country_code]['national_alb'] = int(cell(idx, 'national_alb'))
    dict_2020[country_code]['urban_alb'] = int(cell(idx, 'urban_alb')) #cell(idx, "country_code")
# add data from all these dicts to dict_2015
dict_2020
```

Output:
```
{'AFG': {'country_name': 'Afghanistan',
  'region': 'South Asia',
  'income_level': 'Low income',
  'year': 2020,
  'pop': 38928000,
  'urban_percent': 26,
  'national_alb': 75,
  'urban_alb': 100},
 'ALB': {'country_name': 'Albania',
  'region': 'Europe & Central Asia',
  'income_level': 'Upper middle income',
  'year': 2020,
  'pop': 2878000,
  'urban_percent': 62,
  'national_alb': 95,
  'urban_alb': 96},
 'DZA': {'country_name': 'Algeria',
  'region': 'Middle East & North Africa',
  'income_level': 'Lower middle income',
  'year': 2020,
  'pop': 43851000,
  'urban_percent': 74,
  'national_alb': 94,
  'urban_alb': 96},
 'AND': {'country_name': 'Andorra',
  'region': 'Europe & Central Asia',
  'income_level': 'High income',
  'year': 2020,
  'pop': 77000,
  'urban_percent': 88,
  'national_alb': 100,
  'urban_alb': 100},
 'AGO': {'country_name': 'Angola',
  'region': 'Sub-Saharan Africa',
  'income_level': 'Lower middle income',
  'year': 2020,
  'pop': 32866000,
  'urban_percent': 
```

After you define the data structure `dict_2020`, run the following cell to test whether you have defined it properly.

```python
grader.check("dict_2020")
```

Output:
```
dict_2020 results: All test cases passed!
```

#### From this point onwards, you are only allowed to access data from `water_accessibility.csv` by querying from the **dicts** `dict_2015` and `dict_2020`. You will **lose points** during manual review if you access the data through any other means.

**Question 10:** Output the data from *India* (`country_code`: *IND*) for the `year` *2020*.

Your output **must** be a **dict** mapping each column name to the value for the country *IND* in the year *2020*. You **must** answer this by querying data from `dict_2020`.

The expected output is:
```python
{'country_name': 'India',
 'region': 'South Asia',
 'income_level': 'Lower middle income',
 'year': 2020,
 'pop': 1380004000,
 'urban_percent': 35,
 'national_alb': 90,
 'urban_alb': 94}
```

```python
# compute and store the answer in the variable 'ind_2020_dict', then display it
ind_2020_dict = dict_2020["IND"]
ind_2020_dict
```

Output:
```
{'country_name': 'India',
 'region': 'South Asia',
 'income_level': 'Lower middle income',
 'year': 2020,
 'pop': 1380004000,
 'urban_percent': 35,
 'national_alb': 90,
 'urban_alb': 94}
```

```python
grader.check("q10")
```

Output:
```
q10 results: All test cases passed!
```

**Question 11:** What is the national at least basic (`national_alb`) water supply for *Great Britain* (`country_code`: *GBR*) in the `year` *2015*?

```python
# compute and store the answer in the variable 'gbr_national_alb_2015', then display it
gbr_national_alb_2015 = dict_2015["GBR"]["national_alb"]
gbr_national_alb_2015
```

Output:
```
100
```

```python
grader.check("q11")
```

Output:
```
q11 results: All test cases passed!
```

**Question 12:** How much did the population (`pop`) of *France* **increase** (`country_code`: *FRA*) from the `year` *2015* to *2020*?

```python
# compute and store the answer in the variable 'population_change_fra', then display it
population_change_fra = dict_2020["FRA"]["pop"]
population_change_fra2 = dict_2015['FRA']['pop']
population_change_fra = population_change_fra - population_change_fra2
population_change_fra
```

Output:
```
821000
```

```python
grader.check("q12")
```

Output:
```
q12 results: All test cases passed!
```

**Question 13:** For each `income_level`, find the **total** population (`pop`) of all countries within that `income_level` in *2020*.

Your output **must** be a **dict** where each key is a `income_level`, and the corresponding value is the **sum** of populations (`pop`) of all the countries from that `income_level` in the `year` *2020*.

```python
# compute and store the answer in the variable 'income_level_pops', then display it
income_level_pops = {
    "Low income" : 0,
    'Upper middle income' : 0,
    'Lower middle income' : 0,
    'High income' : 0
}
    

for idx in dict_2020:
    income_types = dict_2020[idx]['income_level']
    population = dict_2020[idx]['pop']

#     if income_types not in income_level_pops:
#         income_level_pops[income_types] = 1
#     else:
#         income_level_pops[income_types] += 1
    income_level_pops[income_types] += population

income_level_pops
```

Output:
```
{'Low income': 514021000,
 'Upper middle income': 2430080000,
 'Lower middle income': 3045857000,
 'High income': 871268000}
```

```python
grader.check("q13")
```

Output:
```
q13 results: All test cases passed!
```

**Question 14:** For each `income_level`, find the **total** population (`pop`) of all countries who have access to at least basic water supply within that `income_level` in *2020*.

Your output **must** be a **dict** where each key is a `income_level`, and the corresponding value is the **sum** of populations (`pop`) which have access to at least basic water supply of all the countries from that `income_level` in the `year` *2020*.

You **must** round the population of **each** country with access to at least basic water supply to the **nearest** integer **before** adding them up.

**Hint:** For each country, the population with at least basic water supply is `pop * national_alb / 100`. 

```python
# compute and store the answer in the variable 'income_level_alb_pops', then display it
# compute and store the answer in the variable 'income_level_alb_pops', then display it
income_level_alb_pops = {
    "Low income" : 0,
    'Upper middle income' : 0,
    'Lower middle income' : 0,
    'High income' : 0
}
value = 0

for idx in dict_2020:
    income_level = dict_2020[idx]['income_level']
    population = dict_2020[idx]['pop']
    nationalalb = dict_2020[idx]['national_alb']
    sum_pop = (population*nationalalb)/100
#     print(population, income_levels)
    
    income_level_alb_pops[income_level] += sum_pop

income_level_alb_pops

```

Output:
```
{'Low income': 302975040.0,
 'Upper middle income': 2321860110.0,
 'Lower middle income': 2687761440.0,
 'High income': 870638000.0}
```

```python
grader.check("q14")
```

Output:
```
q14 results: All test cases passed!
```

**Question 15:** For each `income_level`, find the **percentage** of population (`pop`) of all countries within that `income_level` with at least basic water supply in *2020*.

Your output **must** be a **dict** where each key is a `income_level`, and the corresponding value is the **percentage** of the population (`pop`) which have access to at least basic water supply of all the countries from that `income_level` in the `year` *2020*. The percentages **must** be represented as **int**s between *0* and *100*. You **must** round each of the percentages to the **nearest** integer.

**Hint:** You need to loop through the dictionaries you found in Q13 (or Q14), and for each key, you need to divide the corresponding value in the Q14 dictionary by the value of the same key in the Q13 dictionary and multiply by 100. Take another look at Task 3.6 from lab-p7, if you are not sure how to proceed here.

```python
#compute and store the answer in the variable 'income_level_alb_percent', then display it
income_level_alb_percent = {
    "Low income" : 0,
    'Upper middle income' : 0,
    'Lower middle income' : 0,
    'High income' : 0
}
value = 0

# for idx in dict_2020:
#     income_level = dict_2020[idx]['income_level']
#     population = dict_2020[idx]['pop']
#     nationalalb = dict_2020[idx]['national_alb']
#     sum_pop = (population*nationalalb)/100
#     access_to_water = income_level_alb_pops
#     total_pop = income_level_pops
#     percent = access_to_water / total_pop
    
#     print(population, income_levels)
    
#     income_level_alb_percent[income_level] += percent
    
for idx in income_level_alb_pops:
    income_level_alb_percent[idx] = round(income_level_alb_pops[idx] / income_level_pops[idx] * 100)

income_level_alb_percent = income_level_alb_percent
income_level_alb_percent
```

Output:
```
{'Low income': 59,
 'Upper middle income': 96,
 'Lower middle income': 88,
 'High income': 100}
```

```python
grader.check("q15")
```

Output:
```
q15 results: All test cases passed!
```

### Data Structure 3: Adding `rural_alb`  to `dict_2015` and `dict_2020`

Our dataset has data on the percentage of **national** and **urban** populations with at least basic water supply. However, it is usually **rural** populations which have the greatest difficulty in getting access to water.

Luckily, we are able to calculate **rural_alb** from the given data using the formula:

$$
rural_{alb} = \frac{national_{alb} - \left(urban_{alb} \times \frac{urban\_percent}{100}\right)}{\left(1 - \frac{urban\_percent}{100}\right)}
$$

*If a country has `urban_percent` equal to `100`, then the country has a negligible rural population, and the formula above is not valid. For such countries, we will assume that `rural_alb` is the **same** as `urban_alb`.*

You **must** loop through each country in `dict_2015` and `dict_2020`, and add an **additional** key value pair for each country. The new key should be the string: `"rural_alb"`, and the corresponding value should be the `rural_alb` value for that country as given by the formula above. You **must** round each number to the **nearest** integer.

```python
# add the additional key-value pair to both dicts 'dict_2015' and 'dict_2020' here
# you may display the variable for testing purposes while you define it,
# BUT you MUST remove the line displaying the dicts before submission as the output will be too large to display
#for idx in range(1, len(csv_rows), 2):
 #   part2 = int(cell(idx, 'urban_alb'))*(int(cell(idx, 'urban_percent'))/100)
  #  divide = ((int(cell(idx, 'urban_percent'))/100))
   # rural_alb = (int(cell(idx, "national_alb")) - part2)/divide
```

```python
for key in dict_2020:
    nominator = dict_2020[key]['national_alb'] - (dict_2020[key]['urban_alb']*(dict_2020[key]['urban_percent']/100))
    denominator = 1-(dict_2020[key]['urban_percent']/100)
    if denominator == 0:
        dict_2020[key]['rural_alb'] = dict_2020[key]['urban_alb']
    else:      
        dict_2020[key]['rural_alb'] = round(nominator/denominator)
```

```python
for key in dict_2015:
    nominator = dict_2015[key]['national_alb'] - (dict_2015[key]['urban_alb']*(dict_2015[key]['urban_percent']/100))
    denominator = 1-(dict_2015[key]['urban_percent']/100)
    if denominator == 0:
        dict_2015[key]['rural_alb'] = dict_2015[key]['urban_alb']
    else:      
        dict_2015[key]['rural_alb'] = round(nominator/denominator)
```

**Question 16:** What's the percentage of rural population with at least basic (`rural_alb`) water supply in *Australia* (`country_code`: *AUS*) in *2020*? 

You **must** answer this question by querying data from the dict `dict_2020`.

```python
# compute and store the answer in the variable 'australia_rural_alb_2020', then display it
australia_rural_alb_2020 = dict_2020['AUS']['rural_alb']
australia_rural_alb_2020 = australia_rural_alb_2020
australia_rural_alb_2020
```

Output:
```
100
```

```python
grader.check("q16")
```

Output:
```
q16 results: All test cases passed!
```

### Data Structure 4: `rural_non_alb` bins

We have now managed to extract the percentage of rural population with access to atleast basic water supply for each of the countries in the dataset. We can now use this information to find out the countries whose rural populations do **not** have access to at least basic water supply.

You **must** create two **dict**s (one for the `year` *2015* and one for *2020*) where the keys are the integers *0*, *10*, *20*, ..., *100*. The value corresponding to the integer *0* **must** be a **list** containing the names of all the countries for which their rural population **without** access to at least basic (which we can represent as `rural_non_alb`) water supply is `0 <= rural_non_alb < 10`. Similarly, the value corresponding to the key *10* must be a **list** of all countries for which `10 <= rural_non_alb < 20`, and so on.

**Hints:**
1. You can find `rural_non_alb` as `rural_non_alb = 100 - rural_alb`.
2. You can find the bin which any country falls into by using the formula:
```python
rural_non_alb_bin = ((100 - rural_alb)//10) * 10
```
3. Even if a particular bin has no countries in it, you **must** still create a bin for it in your dict (with the value being an empty list). The starter code below will help you accomplish this.

```python
# compute and store the answer in the variable 'rural_non_alb_bin_2015_dict'

# initialize as an empty dictionary
rural_non_alb_bin_2015_dict = {}

# loop through the keys we want for the dictionary - 0, 10, 20, ..., 100 (inclusive of 100)
# and add them to the dictionary as keys with the value as an empty list
for rural_non_alb_bin in range(11):
    rural_non_alb_bin_2015_dict[rural_non_alb_bin*10] = []
    

for key in dict_2015:
    rural_non_alb_bin = ((100 - dict_2015[key]['rural_alb'])//10) * 10
    rural_non_alb_bin_2015_dict[rural_non_alb_bin].append(dict_2015[key]['country_name'])
# loop through each country and add to the correct bin of rural_non_alb_bin_2015_dict
rural_non_alb_bin_2015_dict
```

Output:
```
{0: ['Andorra',
  'Armenia',
  'Australia',
  'Austria',
  'Bangladesh',
  'Belarus',
  'Belgium',
  'Belize',
  'Bermuda',
  'Bhutan',
  'Bosnia and Herzegovina',
  'Brunei Darussalam',
  'Bulgaria',
  'Canada',
  'Chile',
  'Costa Rica',
  'Cyprus',
  'Czech Republic',
  'Denmark',
  'Estonia',
  'Finland',
  'France',
  'Georgia',
  'Germany',
  'Gibraltar',
  'Greece',
  'Greenland',
  'Guyana',
  'Hungary',
  'Iceland',
  'Ireland',
  'Israel',
  'Jordan',
  'Kazakhstan',
  'Latvia',
  'Lithuania',
  'Luxembourg',
  'Maldives',
  'Malta',
  'Marshall Islands',
  'Mauritius',
  'Monaco',
  'Montenegro',
  'Nauru',
  'Netherlands',
  'New Zealand',
  'North Macedonia',
  'Norway',
  'Palau',
  'Paraguay',
  'Poland',
  'Portugal',
  'Romania',
  'Russian Federation',
  'Samoa',
  'Serbia',
  'Singapore',
  'Spain',
  'Suriname',
  'Sweden',
  'Switzerland',
  'Syrian Arab Republic',
  'Thailand',
  'Tonga',
  'Turkmenistan',
  'Tuvalu',
  'Ukraine',
  'United Kingdom',
  'United Sta
```

```python
# compute and store the answer in the variable 'rural_non_alb_bin_2020_dict'
rural_non_alb_bin_2020_dict = {}

# loop through the keys we want for the dictionary - 0, 10, 20, ..., 100 (inclusive of 100)
# and add them to the dictionary as keys with the value as an empty list
for rural_non_alb_bin in range(11):
    rural_non_alb_bin_2020_dict[rural_non_alb_bin*10] = []
    

for key in dict_2020:
    rural_non_alb_bin = ((100 - dict_2020[key]['rural_alb'])//10) * 10
    rural_non_alb_bin_2020_dict[rural_non_alb_bin].append(dict_2020[key]['country_name'])
# loop through each country and add to the correct bin of rural_non_alb_bin_2015_dict
rural_non_alb_bin_2020_dict
```

Output:
```
{0: ['Albania',
  'Andorra',
  'Armenia',
  'Australia',
  'Austria',
  'Azerbaijan',
  'Bangladesh',
  'Belarus',
  'Belgium',
  'Belize',
  'Bermuda',
  'Bhutan',
  'Bosnia and Herzegovina',
  'Brazil',
  'Brunei Darussalam',
  'Bulgaria',
  'Canada',
  'Chile',
  'Costa Rica',
  'Cuba',
  'Cyprus',
  'Czech Republic',
  'Denmark',
  'Dominican Republic',
  'El Salvador',
  'Estonia',
  'Finland',
  'France',
  'Georgia',
  'Germany',
  'Gibraltar',
  'Greece',
  'Greenland',
  'Guyana',
  'Hungary',
  'Iceland',
  'Iraq',
  'Ireland',
  'Israel',
  'Kazakhstan',
  'Latvia',
  'Lithuania',
  'Luxembourg',
  'Maldives',
  'Malta',
  'Marshall Islands',
  'Mauritius',
  'Mexico',
  'Monaco',
  'Montenegro',
  'Nauru',
  'Netherlands',
  'New Zealand',
  'North Macedonia',
  'Norway',
  'Palau',
  'Paraguay',
  'Philippines',
  'Poland',
  'Portugal',
  'Romania',
  'Russian Federation',
  'Samoa',
  'Serbia',
  'Singapore',
  'Spain',
  'Suriname',
  'Sweden',
  'Switzerland',
  'Syria
```

After you define the data structures `rural_non_alb_bin_2015_dict` and `rural_non_alb_bin_2020_dict`, run the following cell to test whether you have defined them properly.

```python
grader.check("rural_non_alb_bins")
```

Output:
```
rural_non_alb_bins results: All test cases passed!
```

**Question 17:** List all the countries which had `rural_non_alb` value between *0* and *9* (both inclusive) in the `year` *2020*.

You **must** answer this question by querying the the **dict** `rural_non_alb_bin_2020_dict`.

```python
# compute and store the answer in the variable 'bin_0_countries', then display it
bin_0_countries = rural_non_alb_bin_2020_dict[0]
bin_0_countries
```

Output:
```
['Albania',
 'Andorra',
 'Armenia',
 'Australia',
 'Austria',
 'Azerbaijan',
 'Bangladesh',
 'Belarus',
 'Belgium',
 'Belize',
 'Bermuda',
 'Bhutan',
 'Bosnia and Herzegovina',
 'Brazil',
 'Brunei Darussalam',
 'Bulgaria',
 'Canada',
 'Chile',
 'Costa Rica',
 'Cuba',
 'Cyprus',
 'Czech Republic',
 'Denmark',
 'Dominican Republic',
 'El Salvador',
 'Estonia',
 'Finland',
 'France',
 'Georgia',
 'Germany',
 'Gibraltar',
 'Greece',
 'Greenland',
 'Guyana',
 'Hungary',
 'Iceland',
 'Iraq',
 'Ireland',
 'Israel',
 'Kazakhstan',
 'Latvia',
 'Lithuania',
 'Luxembourg',
 'Maldives',
 'Malta',
 'Marshall Islands',
 'Mauritius',
 'Mexico',
 'Monaco',
 'Montenegro',
 'Nauru',
 'Netherlands',
 'New Zealand',
 'North Macedonia',
 'Norway',
 'Palau',
 'Paraguay',
 'Philippines',
 'Poland',
 'Portugal',
 'Romania',
 'Russian Federation',
 'Samoa',
 'Serbia',
 'Singapore',
 'Spain',
 'Suriname',
 'Sweden',
 'Switzerland',
 'Syrian Arab Republic',
 'Thailand',
 'Tonga',
 'Tunisia',
 'Turkmenistan',
 'T
```

```python
grader.check("q17")
```

Output:
```
q17 results: All test cases passed!
```

**Question 18:** What are the countries in the **last** non-empty bin in the `year` *2015*?

Your output **must** be a **list** of the countries in the bin with the **highest** percentage of rural population without at least basic access to water.

**Hint:** You must first find the largest key with a non-empty bin, and then find the value of that key.

```python
# compute and store the answer in the variable 'last_non_empty_bin_2015', then display it

last_non_empty_bin_2015 = []
for idx in rural_non_alb_bin_2015_dict:
    if len(rural_non_alb_bin_2015_dict[idx]) != 0:
        last_non_empty_bin_2015 = rural_non_alb_bin_2015_dict[idx]
    
last_non_empty_bin_2015
```

Output:
```
['Angola', 'Somalia']
```

```python
grader.check("q18")
```

Output:
```
q18 results: All test cases passed!
```

**Question 19:** What countries have **regressed** by moving to a **higher** bin from *2015* to *2020*?

Your answer **must** be a **list** of countries which have regressed by having their percentage of rural population without at least basic access to water move to a bin with a **higher** key.

**Hint:** You may need a *nested* loop to go through all possible combinations of keys in both the dicts `rural_non_alb_bin_2015_dict` and `rural_non_alb_bin_2020_dict`.

```python
# compute and store the answer in the variable 'countries_regressed', then display it
countries_regressed = []
for idx_15 in rural_non_alb_bin_2015_dict:
    for idx_20 in rural_non_alb_bin_2020_dict:
        if idx_20 > idx_15:
            for country in rural_non_alb_bin_2020_dict[idx_20]:
                if country in rural_non_alb_bin_2015_dict[idx_15]:
                    countries_regressed.append(country)

countries_regressed
```

Output:
```
['Jordan',
 'Solomon Islands',
 'Guinea',
 'Zimbabwe',
 'Gabon',
 'Central African Republic']
```

```python
grader.check("q19")
```

Output:
```
q19 results: All test cases passed!
```

**Question 20:** What countries have **improved** by moving to a **lower** bin from *2015* to *2020*?

Your answer **must** be a **list** of countries which have improved by having their percentage of rural population without at least basic access to water move to a bin with a **lower** key.

```python
# compute and store the answer in the variable 'countries_improved', then display it
countries_improved = []

for idx_15 in rural_non_alb_bin_2015_dict:
    for idx_20 in rural_non_alb_bin_2020_dict:
        if idx_20 < idx_15:
            for country in rural_non_alb_bin_2020_dict[idx_20]:
                if country in rural_non_alb_bin_2015_dict[idx_15]:
                    countries_improved.append(country)
countries_improved
```

Output:
```
['Albania',
 'Azerbaijan',
 'Brazil',
 'Cuba',
 'Dominican Republic',
 'El Salvador',
 'Iraq',
 'Mexico',
 'Philippines',
 'Tunisia',
 'Vietnam',
 'Cabo Verde',
 'Colombia',
 'South Africa',
 'Botswana',
 'Ghana',
 'Mali',
 'Morocco',
 'Myanmar',
 'Senegal',
 'Tajikistan',
 'Timor-Leste',
 'Afghanistan',
 'Eswatini',
 'Nigeria',
 'Kenya',
 'Mauritania',
 'Sierra Leone',
 'Togo',
 'Mozambique',
 'Uganda',
 'Somalia']
```

```python
grader.check("q20")
```

Output:
```
q20 results: All test cases passed!
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

**SUBMISSION INSTRUCTIONS**: 1. **Save** the notebook file **now (before you run the next few cells of code)**. 2. **Upload** the zipfile to Gradescope. 3. Check **Gradescope otter** results as soon as the auto-grader execution gets completed. Don't worry about the score showing up as -/100.0. You only need to check that the test cases passed.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for dict_2015 (points: 0)
```
>>> p7_test.check("qdict_2015_test", dict_2015)
True
```


Tests for dict_2020 (points: 0)
```
>>> p7_test.check("qdict_2020_test", dict_2020)
True
```


Tests for get_col_dict (points: 0)
```
>>> test1_output = p7_test.check("qget_col_dict_test1", get_col_dict_test1)
>>> test2_output = p7_test.check("qget_col_dict_test2", get_col_dict_test2)
>>> test3_output = p7_test.check("qget_col_dict_test3", get_col_dict_test3)
>>> test_output = True
>>> 
>>> if test1_output != True:
...     test_output = "get_col_dict_test1 >> " + test1_output
... elif test2_output != True:
...     test_output = "get_col_dict_test2 >> " + test2_output
... elif test3_output != True:
...     test_output = "get_col_dict_test3 >> " + test3_output
>>> 
>>> test_output
True
```


Tests for q1 (points: 5)
```
>>> p7_test.check("q1", highest_pop_country)
True
```


Tests for q10 (points: 5)
```
>>> p7_test.check("q10", ind_2020_dict)
True
```


Tests for q11 (points: 5)
```
>>> p7_test.check("q11", gbr_national_alb_2015)
True
```


Tests for q12 (points: 5)
```
>>> p7_test.check("q12", population_change_fra)
True
```


Tests for q13 (points: 5)
```
>>> p7_test.check("q13", income_level_pops)
True
```


Tests for q14 (points: 5)
```
>>> p7_test.check("q14", income_level_alb_pops)
True
```


Tests for q15 (points: 5)
```
>>> p7_test.check("q15", income_level_alb_percent)
True
```


Tests for q16 (points: 5)
```
>>> p7_test.check("q16", australia_rural_alb_2020)
True
```


Tests for q17 (points: 5)
```
>>> p7_test.check("q17", bin_0_countries)
True
```


Tests for q18 (points: 5)
```
>>> p7_test.check("q18", last_non_empty_bin_2015)
True
```


Tests for q19 (points: 5)
```
>>> p7_test.check("q19", countries_regressed)
True
```


Tests for q2 (points: 5)
```
>>> p7_test.check("q2", highest_pop_inc_country)
True
```


Tests for q20 (points: 5)
```
>>> p7_test.check("q20", countries_improved)
True
```


Tests for q3 (points: 5)
```
>>> p7_test.check("q3", highest_nat_alb_inc_country)
True
```


Tests for q4 (points: 5)
```
>>> p7_test.check("q4", highest_alb_inc_income_level)
True
```


Tests for q5 (points: 5)
```
>>> p7_test.check("q5", total_pop_2020)
True
```


Tests for q6 (points: 5)
```
>>> p7_test.check("q6", urban_pop_percent_2020)
True
```


Tests for q7 (points: 5)
```
>>> p7_test.check("q7", high_income_pop)
True
```


Tests for q8 (points: 5)
```
>>> p7_test.check("q8", least_pop_income_group)
True
```


Tests for q9 (points: 5)
```
>>> p7_test.check("q9", na_max_alb_countries)
True
```


Tests for rural_non_alb_bins (points: 0)
```
>>> rural_non_alb_2015_test = p7_test.check("qrural_non_alb_bin_2015", rural_non_alb_bin_2015_dict)
>>> rural_non_alb_2020_test = p7_test.check("qrural_non_alb_bin_2020", rural_non_alb_bin_2020_dict)
>>> rural_non_alb_bins_test = True
>>> 
>>> if rural_non_alb_2015_test != True:
...     test_output = "rural_non_alb_bin_2015_dict >> " + rural_non_alb_2015_test
... elif rural_non_alb_2020_test != True:
...     test_output = "rural_non_alb_bin_2020_dict >> " + rural_non_alb_2020_test
...     
>>> rural_non_alb_bins_test
True
```
